// Welcome to Starklings: Cairo v1
// This exercise doesn't do anything yet but it still compiles! Cairo file getting run
// needs to have a `main` function. So this file is a valid Cairo file.
// Other exercises will require you to write Cairo code to make the exercise file compile.
// Remove "I AM NOT DONE" comment when you are ready to move to the next exercise.

// use starknet::ContractAddress;
// use traits::TryInto; // Base TryInto trait
// use starknet::Felt252TryIntoContractAddress; // felt > ContractAddress impl
// use starknet::contract_address_to_felt252;
// use option::OptionTrait; // To unwrap

// use debug::PrintTrait; // Just for printing

// #[contract]
// mod testCon {

//     struct Storage{
//         test: LegacyMap::<felt252,felt252>
//     }

//     #[view]
//     fn testFn() -> felt252 {
//         let temp = test::read('a');
//         temp
//     }
// }

fn main() {}